# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    x = str(x)
    y = str(y)
    if len(x)==len(y):
        digx = []
        multx = []
        for i in range(len(x)):
            if x[i] in y:
                if not x[i] in digx:
                    digx.append(x[i])
                    multx.append(1)
                elif x[i] in digx:
                    for j in range(len(digx)):
                        if digx[j]==x[i]:
                            multx[j]+=1
        digy = []
        multy = []
        for k in range(len(x)):
            if y[k] in x:
                if not y[k] in digy:
                    digy.append(y[k])
                    multy.append(1)
                elif y[k] in digy:
                    for l in range(len(digy)):
                        if digy[l]==y[k]:
                            multy[l]+=1
                else:
                    return False
            else:
                return False
        xx = []
        yy = []
        for a in range(len(digx)):
            xx.append((digx[a],multx[a]))
            yy.append((digy[a],multy[a]))
        dx = dict(xx)
        dy = dict(yy)
        if dx==dy:
            return True   
    else:
        return False
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
two = []
r = 0
while r<1000000:
    if same_digits(r, 2*r):
        two = [r]
        r+=1
    else:
        r+=1
two
```

Output:
```
[428571]
```

```python
three = []
r = 0
while r<100:
    if same_digits(r, 3*r):
        two = [r]
        r+=1
    else:
        r+=1
three
```

Output:
```
[]
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```