# Algorithms Exercise 2

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import seaborn as sns
import numpy as np
```

## Peak finding

Write a function `find_peaks` that finds and returns the indices of the local maxima in a sequence. Your function should:

* Properly handle local maxima at the endpoints of the input array.
* Return a Numpy array of integer indices.
* Handle any Python iterable as input.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def find_peaks(a):
    """Find the indices of the local maxima in a sequence."""
    peaks = []
    for i in range(len(a)):
        if a[i]>a[i-1] and a[i]==len(a)-1:
            peaks.append(i)
        elif a[i]!=len(a)-1:
            if a[i]>a[i-1] and a[i]>a[i+1] or a[i-1]==a[len(a)] and a[i]>a[i+1]:
                peaks.append(i)
```

```python
a = [2,0,1,0,2,0,1]
p1 = find_peaks(a)
print(p1)
```

Output:
```
<ipython-input-32-3c7c2266d67e> in <module>()
      1 a = [2,0,1,0,2,0,1]
----> 2 p1 = find_peaks(a)
      3 print(p1)

<ipython-input-31-fd06c8192dc2> in find_peaks(a)
      6             peaks.append(i)
      7         elif a[i]!=len(a)-1:
----> 8             if a[i]>a[i-1] and a[i]>a[i+1] or a[i-1]==a[len(a)] and a[i]>a[i+1]:
      9                 peaks.append(i)

IndexError: list index out of range
```

<!-- grade_id: algorithmsex02a, points: 5, grade: true, solution: false -->
```python
p1 = find_peaks([2,0,1,0,2,0,1])
assert np.allclose(p1, np.array([0,2,4,6]))
p2 = find_peaks(np.array([0,1,2,3]))
assert np.allclose(p2, np.array([3]))
p3 = find_peaks([3,2,1,0])
assert np.allclose(p3, np.array([0]))
```

Output:
```
<ipython-input-17-e9503c7dbdc3> in <module>()
      1 p1 = find_peaks([2,0,1,0,2,0,1])
----> 2 assert np.allclose(p1, np.array([0,2,4,6]))
      3 p2 = find_peaks(np.array([0,1,2,3]))
      4 assert np.allclose(p2, np.array([3]))
      5 p3 = find_peaks([3,2,1,0])

/usr/local/lib/python3.4/dist-packages/numpy/core/numeric.py in allclose(a, b, rtol, atol)
   2217     y = array(y, dtype=dtype, copy=False)
   2218 
-> 2219     xinf = isinf(x)
   2220     yinf = isinf(y)
   2221     if any(xinf) or any(yinf):

TypeError: ufunc 'isinf' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

Here is a string with the first 10000 digits of $\pi$ (after the decimal). Write code to perform the following:

* Convert that string to a Numpy array of integers.
* Find the indices of the local maxima in the digits of $\pi$.
* Use `np.diff` to find the distances between consequtive local maxima.
* Visualize that distribution using an appropriately customized histogram.

```python
from sympy import pi, N
pi_digits_str = str(N(pi, 10001))[2:]
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: algorithmsex02b, points: 5, grade: true, solution: false -->
```python
assert True # use this for grading the pi digits histogram
```