# Algorithms Exercise 1

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

## Word counting

Write a function `tokenize` that takes a string of English text returns a list of words. It should also remove [stop words](http://en.wikipedia.org/wiki/Stop_words), which are common short words that are often removed before natural language processing. Your function should have the following logic:

* Split the string into lines using `splitlines`.
* Split each line into a list of words and merge the lists for each line.
* Use Python's builtin `filter` function to remove all punctuation.
* If `stop_words` is a list, remove all occurences of the words in the list.
* If `stop_words` is a space delimeted string of words, split them and remove them.
* Remove any remaining empty words.
* Make all words lowercase.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def tokenize(s, stop_words=None, punctuation='`~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t'):
    """Split a string into a list of words, removing punctuation and stop words."""
    low = []
    t = s.splitlines()
    for i in range(len(t)):
        u = t[i].split()
        for j in range(len(u)):
            low.append(u[j])                        # Turns multi-lined string into neat list
    no_more_punc = []
    for k in range(len(low)):
        no_more_punc.append(''.join(list(filter(lambda x: not x in punctuation, low[k]))))      # Removes punctuation
    if type(stop_words)==list:
        no_more_stops = list(filter(lambda x: not x in stop_words, no_more_punc))               # Removes stop words
    elif type(stop_words)==str:
        no_more_stops = list(filter(lambda x: not x in stop_words.split(), no_more_punc))       # for different input types
    elif stop_words==None:
        no_more_stops = no_more_punc
    no_gaps = list(filter(lambda x: not x=='', no_more_stops))                                  # Removes empty strings
    all_low = []
    for l in range(len(no_gaps)):
        all_low.append(no_gaps[l].lower())                                                      # Removes any capitalization
    return all_low
```

```python
tokenize("This, is the way; that things will end", stop_words=['the', 'is'])
```

Output:
```
['this', 'way', 'that', 'things', 'will', 'end']
```

<!-- grade_id: algorithmsex01a, points: 2, grade: true, solution: false -->
```python
assert tokenize("This, is the way; that things will end", stop_words=['the', 'is']) == \
    ['this', 'way', 'that', 'things', 'will', 'end']
wasteland = """
APRIL is the cruellest month, breeding
Lilacs out of the dead land, mixing
Memory and desire, stirring
Dull roots with spring rain.
"""

assert tokenize(wasteland, stop_words='is the of and') == \
    ['april','cruellest','month','breeding','lilacs','out','dead','land',
     'mixing','memory','desire','stirring','dull','roots','with','spring',
     'rain']
```

Write a function `count_words` that takes a list of words and returns a dictionary where the keys in the dictionary are the unique words in the list and the values are the word counts.

```python
p = [1,2,3]
q = ['one','two','three']
pq = []
for i in range(len(p)):
    pq.append((p[i],q[i]))
d = dict(pq)
d
```

Output:
```
{1: 'one', 2: 'two', 3: 'three'}
```

<!-- grade_id: algorithmsex01b, points: 2, grade: false, solution: true -->
```python
def count_words(data):
    """Return a word count dictionary from the list of words in data."""
    word = []
    count = []
    for i in range(len(data)):
        if not data[i] in word:
            word.append(data[i])
            count.append(1)
        elif data[i] in word:
            for j in range(len(word)):
                if word[j]==data[i]:
                    count[j]+=1
    wc = []
    for j in range(len(word)):
        wc.append((word[j],count[j]))
    wcd = dict(wc)
    return wcd
```

```python
count_words(tokenize('this and the this from and a a a'))
```

Output:
```
{'a': 3, 'and': 2, 'from': 1, 'the': 1, 'this': 2}
```

<!-- grade_id: algorithmsex01b, points: 2, grade: true, solution: false -->
```python
assert count_words(tokenize('this and the this from and a a a')) == \
    {'a': 3, 'and': 2, 'from': 1, 'the': 1, 'this': 2}
```

Write a function `sort_word_counts` that return a list of sorted word counts:

* Each element of the list should be a `(word, count)` tuple.
* The list should be sorted by the word counts, with the higest counts coming first.
* To perform this sort, look at using the `sorted` function with a custom `key` and `reverse`
  argument.

```python
wiggity = [(1,'one'),(3,'three'),(2,'two')]
sorted(wiggity, key=lambda x: x[0], reverse=True)
```

Output:
```
[(3, 'three'), (2, 'two'), (1, 'one')]
```

```python
boogity = {1:2,2:3,3:4}
list(iter(boogity)),boogity[1]
```

Output:
```
([1, 2, 3], 2)
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def sort_word_counts(wc):
    """Return a list of 2-tuples of (word, count), sorted by count descending."""
    word = list(iter(wc))
    count = []
    for i in word:
        count.append(wc[i])
    tups = []
    for j in range(len(word)):
        tups.append((word[j],count[j]))
    return sorted(tups, key=lambda x: x[1], reverse=True)
```

```python
sort_word_counts(count_words(tokenize('this and a the this this and a a a')))
```

Output:
```
[('a', 4), ('this', 3), ('and', 2), ('the', 1)]
```

<!-- grade_id: algorithmsex01c, points: 2, grade: true, solution: false -->
```python
assert sort_word_counts(count_words(tokenize('this and a the this this and a a a'))) == \
    [('a', 4), ('this', 3), ('and', 2), ('the', 1)]
```

Perform a word count analysis on Chapter 1 of Moby Dick, whose text can be found in the file `mobydick_chapter1.txt`:

* Read the file into a string.
* Tokenize with stop words of `'the of and a to in is it that as'`.
* Perform a word count, the sort and save the result in a variable named `swc`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
with open('mobydick_chapter1.txt', 'r') as f:
    read_text = f.read()
f.closed
```

Output:
```
True
```

```python
swc = sort_word_counts(count_words(tokenize(read_text, 'the of and a to in is it that as')))
```

<!-- grade_id: algorithmsex01d, points: 2, grade: true, solution: false -->
```python
assert swc[0]==('i',43)
assert len(swc)==848
```

Create a "Cleveland Style" [dotplot](http://en.wikipedia.org/wiki/Dot_plot_%28statistics%29) of the counts of the top 50 words using Matplotlib. If you don't know what a dotplot is, you will have to do some research...

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: algorithsex01e, points: 2, grade: true, solution: false -->
```python
assert True # use this for grading the dotplot
```