# Integration Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

## 2d polar integration

The 2d polar integral of a scalar function $f(r, \theta)$ is defined as:

$$ I(r_{max}) = \int_0^{r_{max}} \int_0^{2\pi} f(r, \theta) r dr d\theta $$

Write a function `integrate_polar(f, rmax)` that performs this integral numerically using `scipy.integrate.dblquad`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integrate_polar(f, rmax):
    """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
    # y=r, x=t, f(y,x) = f(r,t)
    F = lambda r,t: f(r,t)*r
    I, e = integrate.dblquad(F, 0, 2*np.pi, lambda x: 0, lambda x: rmax)
    return I
```

```python
integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf)
```

Output:
```
3.141592653589792
```

<!-- grade_id: integrationex04a, points: 10, grade: true, solution: false -->
```python
assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)
```